# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string ="""Amy is 5 years old, and her sister Mary is 2 years old.Ruth and Peter, their parents, have 3 kids."""
    pattern='[A-Z][a-z]*'
    result=re.findall(pattern,simple_string)
    return result
    #raise NotImplementedError()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()
    result = re.findall('(.*): B',grades)
    return result
    #raise NotImplementedError()
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

Output:
```
/var/folders/m5/gjzyxs7575z5k5j0xxzrvzlc0000gn/T/ipykernel_9841/1885468736.py in <module>
----> 1 assert len(grades()) == 16

/var/folders/m5/gjzyxs7575z5k5j0xxzrvzlc0000gn/T/ipykernel_9841/4293573819.py in grades()
      1 import re
      2 def grades():
----> 3     with open ("assets/grades.txt", "r") as file:
      4         grades = file.read()
      5     result = re.findall('(.*): B',grades)

FileNotFoundError: [Errno 2] No such file or directory: 'assets/grades.txt'
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    result=list()
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
        pattern="""
        (?P<host>[0-9]+\.[0-9]+\.[0-9]+\.[0-9]+)
        (\ - \ )
        (?P<user_name>(\w*)(\S))
        (\  \S)
        (?P<time>\d+\S\w*\S\d+\S\d+\S\d+\S\d+\s\S\d+)
        (\S\s\S)
        (?P<request>\w*\s\S*\s\w*\S\d.\d*)
        """
    for item in re.finditer(pattern,logdata,re.VERBOSE):
        result.append(item.groupdict())
    return result
    #raise NotImplementedError()
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

Output:
```
/var/folders/m5/gjzyxs7575z5k5j0xxzrvzlc0000gn/T/ipykernel_9841/3220867527.py in <module>
----> 1 assert len(logs()) == 979
      2 
      3 one_item={'host': '146.204.224.152',
      4   'user_name': 'feest6811',
      5   'time': '21/Jun/2019:15:45:24 -0700',

/var/folders/m5/gjzyxs7575z5k5j0xxzrvzlc0000gn/T/ipykernel_9841/2475100478.py in logs()
      2 def logs():
      3     result=list()
----> 4     with open("assets/logdata.txt", "r") as file:
      5         logdata = file.read()
      6         pattern="""

FileNotFoundError: [Errno 2] No such file or directory: 'assets/logdata.txt'
```

# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```



```python
def proportion_of_education():
    # your code goes here
    import pandas as pd
    df = pd.read_csv('assets/NISPUF17.csv', index_col=0)
    EDUC1 = df['EDUC1'].value_counts(normalize=True)
    EDUC1_dict = {"less than high school": EDUC1.iloc[3],
                   "high school": EDUC1.iloc[2],
                   "more than high school but not college": EDUC1.iloc[1],
                   "college": EDUC1.iloc[0]}
    return EDUC1_dict
    # YOUR CODE HERE
    #raise NotImplementedError()
```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
def average_influenza_doses():
    # YOUR CODE HERE
    import pandas as pd
    df = pd.read_csv('assets/NISPUF17.csv')
    df1 = df[df["CBF_01"] == 1]
    df2 = df[df["CBF_01"] == 2]
    
    return (df1["P_NUMFLU"].mean(),df2["P_NUMFLU"].mean())
    #raise NotImplementedError()
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    import pandas as pd
    df = pd.read_csv('assets/NISPUF17.csv')
    hdyp = df[(df["HAD_CPOX"] ==1) & (df["P_NUMVRC"] >= 1)]
    hdnp = df[(df["HAD_CPOX"] ==2) & (df["P_NUMVRC"] >= 1)]

    m_ratio_p = len(hdyp[hdyp["SEX"] == 1])/len(hdnp[hdnp["SEX"] == 1])
    f_ratio_p = len(hdyp[hdyp["SEX"] == 2])/len(hdnp[hdnp["SEX"] == 2])

    
    result_dic ={"male": m_ratio_p,
                "female": f_ratio_p}
    return result_dic
    #raise NotImplementedError()
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    df = pd.read_csv('assets/NISPUF17.csv', index_col=0)
    df = df[df["HAD_CPOX"] <= 2]
    df = df[~df["P_NUMVRC"].isna() & ~df["HAD_CPOX"].isna()]

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["HAD_CPOX"], df["P_NUMVRC"])
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
#     raise NotImplementedError()
```

# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 3-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

```python
def answer_one():
    # YOUR CODE HERE
    
    # Energy dataset
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows=0,header=17,skipfooter=38)
    del Energy['Unnamed: 0']
    del Energy['Unnamed: 2']
    Energy.rename(columns={'Petajoules':'Energy Supply',
                               'Gigajoules':'Energy Supply per Capita'
                               ,'%':'% Renewable',
                                'Unnamed: 1':'Country'},inplace=True)
    # change name Energy
    Energy = Energy.replace({"...": np.NaN,"Republic of Korea": "South Korea",
                             "United States of America": "United States",
                            "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                            "China, Hong Kong Special Administrative Region": "Hong Kong"})
    # conver Energy Supply to Gigajoules
    Energy['Energy Supply'] *= 1000000
    # delete (...)
    Energy["Country"] = Energy["Country"].str.replace(' [()].*[)]', '',regex=True)
    
    # GDP dataset
    GDP = pd.read_csv('assets/world_bank.csv',header=4)
    # change name GDP
    GDP = GDP.replace({"Korea, Rep.": "South Korea",
                   "Iran, Islamic Rep.": "Iran",
                   "Hong Kong SAR, China": "Hong Kong"})
    GDP = GDP.rename(columns={"Country Name":"Country"})
    
    # ScimEn dataset
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    
    # setting index to Country
    Energy.set_index(['Country'],inplace=True) 
    ScimEn.set_index(['Country'],inplace=True)
    GDP.set_index(['Country'],inplace=True)
    
    # define ScimEn dataset for merge
    df1 = ScimEn[['Rank', 'Documents', 'Citable documents', 'Citations',
              'Self-citations', 'Citations per document', 'H index']]
    df1 = df1[df1['Rank'] <= 15]
    
    # define Energy dataset for merge
    df2 = Energy[['Energy Supply', 'Energy Supply per Capita', '% Renewable']]
    
    # define GDP dataset for merge
    df3 = GDP[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    
    # define final dataset 
    df_final = df1.merge(df2,left_index=True,
                     right_index=True,how='inner').merge(df3,left_index=True,
                                                         right_index=True,how='inner')
    return df_final
    #raise NotImplementedError()
```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

```python
def answer_two():
    # YOUR CODE HERE
    # Energy dataset
    Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows=0,header=17,skipfooter=38)
    del Energy['Unnamed: 0']
    del Energy['Unnamed: 2']
    Energy.rename(columns={'Petajoules':'Energy Supply',
                               'Gigajoules':'Energy Supply per Capita'
                               ,'%':'% Renewable',
                                'Unnamed: 1':'Country'},inplace=True)
    # change name Energy
    Energy = Energy.replace({"...": np.NaN,"Republic of Korea": "South Korea",
                             "United States of America": "United States",
                            "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                            "China, Hong Kong Special Administrative Region": "Hong Kong"})
    # conver Energy Supply to Gigajoules
    Energy['Energy Supply'] *= 1000000
    # delete (...)
    Energy["Country"] = Energy["Country"].str.replace(' [()].*[)]', '',regex=True)
    
    # GDP dataset
    GDP = pd.read_csv('assets/world_bank.csv',header=4)
    # change name GDP
    GDP = GDP.replace({"Korea, Rep.": "South Korea",
                   "Iran, Islamic Rep.": "Iran",
                   "Hong Kong SAR, China": "Hong Kong"})
    GDP = GDP.rename(columns={"Country Name":"Country"})
    
    # ScimEn dataset
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    
    # setting index to Country
    Energy.set_index(['Country'],inplace=True) 
    ScimEn.set_index(['Country'],inplace=True)
    GDP.set_index(['Country'],inplace=True)
    
    # define ScimEn dataset for merge
    df1 = ScimEn[['Rank', 'Documents', 'Citable documents', 'Citations',
              'Self-citations', 'Citations per document', 'H index']]
    df1 = df1[df1['Rank'] <= 15]
    
    # define Energy dataset for merge
    df2 = Energy[['Energy Supply', 'Energy Supply per Capita', '% Renewable']]
    
    # define GDP dataset for merge
    df3 = GDP[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    
    # define final dataset 
#     df_final_1 = df1.merge(df2,left_index=True,
#                      right_index=True,how='outer').merge(df3,left_index=True,
#                                                          right_index=True,how='outer')
#     df_final_2 = df1.merge(df2,left_index=True,
#                      right_index=True,how='inner').merge(df3,left_index=True,
#                                                          right_index=True,how='inner')
#     delta = df_final_1.shape[0] - df_final_2.shape[0]
#     return delta

    union = pd.merge(pd.merge(Energy, GDP, on='Country', how='outer'), ScimEn, on='Country', how='outer')
    intersect = pd.merge(pd.merge(Energy, GDP, on='Country'), ScimEn, on='Country')
    return len(union)-len(intersect)
    #raise NotImplementedError()
```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

```python
def answer_three():
    # YOUR CODE HERE
    # if there was no other functions
    
#     GDP = pd.read_csv('assets/world_bank.csv',header=4)
#     GDP = GDP.replace({"Korea, Rep.": "South Korea",
#                        "Iran, Islamic Rep.": "Iran",
#                        "Hong Kong SAR, China": "Hong Kong"})
#     GDP = GDP.rename(columns={"Country Name":"Country"})
#     df_avg = GDP[['Country','2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
#     df_avg = df_avg.set_index('Country')
#     df_avg['mean'] = df_avg.mean(axis=1)
#     df_avg = df_avg[['mean']].squeeze()
    
#     return df_avg.nlargest(15)

    # by using other functions
    
    Top15 = answer_one()
    rows = ['2006', '2007', '2008','2009', '2010', '2011', '2012', '2013', '2014', '2015']
    Top15["avgGDP"] = Top15[rows].mean(axis=1)
    return Top15.sort_values("avgGDP",ascending=False)["avgGDP"]
    #raise NotImplementedError()
```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

```python
def answer_four():
    # YOUR CODE HERE
    # if there was no other functions 
    
#     GDP = pd.read_csv('assets/world_bank.csv',header=4)
#     GDP = GDP.replace({"Korea, Rep.": "South Korea",
#                        "Iran, Islamic Rep.": "Iran",
#                        "Hong Kong SAR, China": "Hong Kong"})
#     GDP = GDP.rename(columns={"Country Name":"Country"})
#     df_avg = GDP[['Country','2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
#     df_avg = df_avg.set_index('Country')
#     df_avg['mean'] = df_avg.mean(axis=1)
#     se_avg = df_avg[['mean']].squeeze()
#     se_avg = se_avg.sort_values(ascending=False)
#     country_1 = se_avg_sorted[se_avg_sorted == se_avg_sorted[5]].index[0]
#     return (df_avg.loc[country_1][0] - df_avg.loc[country_1][-2])

    # by using other functions
    Top15 = answer_one()
    Top15["AvgGDP"] = answer_three()
    Top15.sort_values("AvgGDP", ascending=False, inplace=True)
    final = Top15.iloc[5]['2015']
    initial = Top15.iloc[5]['2006']
    return abs(final - initial)
    #raise NotImplementedError()
```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

```python
def answer_five():
    # YOUR CODE HERE
    # if there was no other functions 
    
#     Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows=0,header=17,skipfooter=38)
#     del Energy['Unnamed: 0']
#     del Energy['Unnamed: 2']
#     Energy.rename(columns={'Petajoules':'Energy Supply',
#                                'Gigajoules':'Energy Supply per Capita'
#                                ,'%':'% Renewable',
#                                 'Unnamed: 1':'Country'},inplace=True)
#     # change name Energy
#     Energy = Energy.replace({"...": np.NaN})
#     # conver Energy Supply to Gigajoules
    
#     Energy['Energy Supply'] *= 1000000
#     return Energy['Energy Supply'].mean()

    # by using other functions
    Top15 = answer_one()
    return Top15['Energy Supply per Capita'].mean()
    #raise NotImplementedError()
```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

```python
def answer_six():
    # YOUR CODE HERE
    # if there was no other functions 
    
#     Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows=0,header=17,skipfooter=38)
#     del Energy['Unnamed: 0']
#     del Energy['Unnamed: 2']
#     Energy.rename(columns={'Petajoules':'Energy Supply',
#                                'Gigajoules':'Energy Supply per Capita'
#                                ,'%':'% Renewable',
#                                 'Unnamed: 1':'Country'},inplace=True)
#     # change name Energy
#     Energy = Energy.replace({"...": np.NaN})


#     max_ren = Energy['% Renewable'].max()
#     max_country = Energy['Country'][Energy['% Renewable'][Energy['% Renewable'] == Energy['% Renewable'].max()].index[0]]

#     return (max_country,max_ren)

    # by using other functions
    
    Top15 = answer_one()
    return Top15['% Renewable'].argmax(), Top15['% Renewable'].max()
    #raise NotImplementedError()
```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

```python
def answer_seven():
    # YOUR CODE HERE
    # if there was no other functions 
    
#     ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')

#     ScimEn['RST'] = (ScimEn['Self-citations']/ScimEn['Citable documents'] 
#                                               

#     max_ratio = ScimEn['RST'].max()
#     max_country = ScimEn['Country'][ScimEn['RST'][ScimEn['RST']== ScimEn['RST'].max()].index[0]]

#     return (max_country,max_ratio)

    # by using other functions
    
    Top15 = answer_one()
    Top15["Ratio"] = Top15["Self-citations"] / Top15['Citations']
    return (Top15["Ratio"].argmax(), Top15["Ratio"].max())
    #raise NotImplementedError()
```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

```python
def answer_eight():
    # YOUR CODE HERE
    # if there was no other functions
#     ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
#     Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows=0,header=17,skipfooter=38)
#     del Energy['Unnamed: 0']
#     del Energy['Unnamed: 2']
#     Energy.rename(columns={'Petajoules':'Energy Supply',
#                                'Gigajoules':'Energy Supply per Capita'
#                                ,'%':'% Renewable',
#                                 'Unnamed: 1':'Country'},inplace=True)
#     # change name Energy
#     Energy = Energy.replace({"...": np.NaN})
#     # conver Energy Supply to Gigajoules
#     Energy['Energy Supply'] *= 1000000
#     Energy["Population"] =  Energy['Energy Supply'] / Energy['Energy Supply per Capita']
#     final = Energy.sort_values("Population", ascending=False)
#     return final.iloc[2].name

    # by using other functions
    
    Top15 = answer_one()
    Top15["Population"] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    final = Top15.sort_values("Population", ascending=False)
    return final.iloc[2].name
    #raise NotImplementedError()
```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

```python
def answer_nine():
    # YOUR CODE HERE
    # if there was no other functions
    
#     Energy = pd.read_excel('assets/Energy Indicators.xls',skiprows=0,header=17,skipfooter=38)
#     del Energy['Unnamed: 0']
#     del Energy['Unnamed: 2']
#     Energy.rename(columns={'Petajoules':'Energy Supply',
#                                'Gigajoules':'Energy Supply per Capita'
#                                ,'%':'% Renewable',
#                                 'Unnamed: 1':'Country'},inplace=True)
#     # change name Energy
#     Energy = Energy.replace({"...": np.NaN})
#     # conver Energy Supply to Gigajoules
#     Energy['Energy Supply'] *= 1000000
#     Energy["Population"] =  Energy['Energy Supply'] / Energy['Energy Supply per Capita']
#     ScimEn["Citable docs per Capita"] = ScimEn["Citable documents"] / Energy["Population"]
#     ScimEn["Citable docs per Capita"].corr(Energy['Energy Supply per Capita']) 

    # by using other functions
    
    Top15 = answer_one()
    Top15["Population"] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15["Citable docs per Capita"] = Top15["Citable documents"] / Top15["Population"]
    return Top15["Citable docs per Capita"].corr(Top15['Energy Supply per Capita'])
    #raise NotImplementedError()
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

```python
def answer_ten():
    # YOUR CODE HERE
    Top15 = answer_one()
    reference = Top15["% Renewable"].median(axis=0)
    Top15["HighRenew"] = Top15.apply(lambda x: 1 if x["% Renewable"] >= reference else 0, axis=1)
    Top15.sort_values(by='Rank', inplace=True)
    return Top15["HighRenew"] 
    #raise NotImplementedError()
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

```python
def answer_eleven():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    
    Top15 = answer_one()
    final_c = pd.DataFrame(columns=['size', 'sum', 'mean', 'std'])
    Top15['Estimate Population'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    for group, contient in Top15.groupby(ContinentDict):
        final_c.loc[group] = [len(contient), contient['Estimate Population'].sum(),
                            contient['Estimate Population'].mean(),
                            contient['Estimate Population'].std()]
    return final_c
    #raise NotImplementedError()
```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

```python
def answer_twelve():
    Top15 = answer_one()
    ContinentDict  = {'China':'Asia',
                      'United States':'North America', 
                      'Japan':'Asia', 
                      'United Kingdom':'Europe', 
                      'Russian Federation':'Europe', 
                      'Canada':'North America', 
                      'Germany':'Europe', 
                      'India':'Asia',
                      'France':'Europe', 
                      'South Korea':'Asia', 
                      'Italy':'Europe', 
                      'Spain':'Europe', 
                      'Iran':'Asia',
                      'Australia':'Australia', 
                      'Brazil':'South America'
                     }
    
    Top15['Continent'] = Top15.index.map(lambda c: ContinentDict[c])
    
    # https://pandas.pydata.org/pandas-docs/stable/generated/pandas.cut.html
    Top15['% Renewable bins'] = pd.cut(Top15['% Renewable'], 5)
    
    return Top15.groupby(['Continent','% Renewable bins']).size()


    #raise NotImplementedError()


```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

```python
def answer_thirteen():
    # YOUR CODE HERE
    Top15 = answer_one()
    Top15["Population"] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    return Top15['Population'].apply(lambda x: '{0:,}'.format(x))
    #raise NotImplementedError()
```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
    
plot_optional()
```

# Assignment 4
## Description
In this assignment you must read in a file of metropolitan regions and associated sports teams from [assets/wikipedia_data.html](assets/wikipedia_data.html) and answer some questions about each metropolitan region. Each of these regions may have one or more teams from the "Big 4": NFL (football, in [assets/nfl.csv](assets/nfl.csv)), MLB (baseball, in [assets/mlb.csv](assets/mlb.csv)), NBA (basketball, in [assets/nba.csv](assets/nba.csv) or NHL (hockey, in [assets/nhl.csv](assets/nhl.csv)). Please keep in mind that all questions are from the perspective of the metropolitan region, and that this file is the "source of authority" for the location of a given sports team. Thus teams which are commonly known by a different area (e.g. "Oakland Raiders") need to be mapped into the metropolitan region given (e.g. San Francisco Bay Area). This will require some human data understanding outside of the data you've been given (e.g. you will have to hand-code some names, and might need to google to find out where teams are)!

For each sport I would like you to answer the question: **what is the win/loss ratio's correlation with the population of the city it is in?** Win/Loss ratio refers to the number of wins over the number of wins plus the number of losses. Remember that to calculate the correlation with [`pearsonr`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.pearsonr.html), so you are going to send in two ordered lists of values, the populations from the wikipedia_data.html file and the win/loss ratio for a given sport in the same order. Average the win/loss ratios for those cities which have multiple teams of a single sport. Each sport is worth an equal amount in this assignment (20%\*4=80%) of the grade for this assignment. You should only use data **from year 2018** for your analysis -- this is important!

## Notes

1. Do not include data about the MLS or CFL in any of the work you are doing, we're only interested in the Big 4 in this assignment.
2. I highly suggest that you first tackle the four correlation questions in order, as they are all similar and worth the majority of grades for this assignment. This is by design!
3. It's fair game to talk with peers about high level strategy as well as the relationship between metropolitan areas and sports teams. However, do not post code solving aspects of the assignment (including such as dictionaries mapping areas to teams, or regexes which will clean up names).
4. There may be more teams than the assert statements test, remember to collapse multiple teams in one city into a single value!

## Question 1
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **NHL** using **2018** data.

```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

nhl_df=pd.read_csv("assets/nhl.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def clean_nhl_df():
    # load data
    nhl_df=pd.read_csv("assets/nhl.csv")
    cities=pd.read_html("assets/wikipedia_data.html")[1]
    cities=cities.iloc[:-1,[0,3,5,6,7,8]]

    # cleaning the cities dataframe
    cities["NHL"] = cities["NHL"].apply(lambda x: re.sub(r"\[.+\]", "", x))
    cities["NHL"] = cities["NHL"].replace({"RangersIslandersDevils": "Rangers,Islanders,Devils",
                                           "KingsDucks": "Kings,Ducks",
                                           "Red Wings": "Red,Wings", 
                                           "Maple Leafs": "Maple,Leafs", 
                                           "Blue Jackets": "Blue,Jackets",
                                           "Golden Knights": "Golden,Knights" })
    cities["NHL"] = cities["NHL"].apply(lambda x: x.split(","))
    cities = cities.explode("NHL")

    # cleaning the nhl_df dataframe
    nhl_df = nhl_df[nhl_df["year"] == 2018]
    nhl_df["team"] = nhl_df["team"].apply(lambda x: x.replace("*", ""))
    nhl_df["team"] = nhl_df["team"].apply(lambda x: x.split(" ")[-1])

    # merge the dataframes
    df = pd.merge(cities, nhl_df, left_on="NHL", right_on="team")
    df = df[["Metropolitan area", "Population (2016 est.)[8]", "NHL", "team", "W", "L"]]
    df["W-L%"] = df["W"].astype("int")/(df["W"].astype("int") + df["L"].astype("int"))
    df["Population (2016 est.)[8]"] = df["Population (2016 est.)[8]"].astype("float")
    df["W-L%"] = df["W-L%"].astype("float")

    # drop duplicated columns
    df.loc[df["Metropolitan area"] == "New York City", "W-L%"] = df.loc[df["Metropolitan area"] == "New York City", "W-L%"].mean() # mean of NY W-L%
    df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"] = df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"].mean() # mean of LA W-L%
    df = df.drop_duplicates(subset="Metropolitan area").reset_index()
    df = df.drop(columns="index")
    
    return df

def nhl_correlation(): 
    # YOUR CODE HERE
#     raise NotImplementedError()
    df = clean_nhl_df()

    population_by_region = df["Population (2016 est.)[8]"] # pass in metropolitan area population from cities
    win_loss_by_region = df["W-L%"] # pass in win/loss ratio from nhl_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q1: Your lists must be the same length"
    assert len(population_by_region) == 28, "Q1: There should be 28 teams being analysed for NHL"
    
    result = stats.pearsonr(population_by_region, win_loss_by_region)
    
    return stats.pearsonr(population_by_region, win_loss_by_region)
```

## Question 2
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **NBA** using **2018** data.

```python
nba_df=pd.read_csv("assets/nba.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def clean_nba_df():
    # load data
    nba_df=pd.read_csv("assets/nba.csv")
    cities=pd.read_html("assets/wikipedia_data.html")[1]
    cities=cities.iloc[:-1,[0,3,5,6,7,8]]
    
    # cleaning the cities dataframe
    cities["NBA"] = cities["NBA"].apply(lambda x: re.sub(r"\[.+\]", "", x))
    cities["NBA"] = cities["NBA"].replace({"KnicksNets": "Knicks,Nets",
                                           "LakersClippers": "Lakers,Clippers",
                                           "Trail Blazers": "Trail,Blazers"})
    cities["NBA"] = cities["NBA"].apply(lambda x: x.split(","))
    cities = cities.explode("NBA")

    # cleaning the nhl_df dataframe
    nba_df = nba_df[nba_df["year"] == 2018]
    nba_df["team"] = nba_df["team"].apply(lambda x: re.sub(r"(\*)*\s\(\d+\)", "", x))
    nba_df["team"] = nba_df["team"].apply(lambda x: x.split(" ")[-1])

    # merge the dataframes
    df = pd.merge(cities, nba_df, left_on="NBA", right_on="team")
    df = df.rename(columns={"W/L%": "W-L%"})
    df = df[["Metropolitan area", "Population (2016 est.)[8]", "NBA", "team", "W", "L", "W-L%"]]
    df["Population (2016 est.)[8]"] = df["Population (2016 est.)[8]"].astype("float")
    df["W-L%"] = df["W-L%"].astype("float")
    
    # drop duplicated columns
    df.loc[df["Metropolitan area"] == "New York City", "W-L%"] = df.loc[df["Metropolitan area"] == "New York City", "W-L%"].mean() # mean of NY W-L%
    df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"] = df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"].mean() # mean of LA W-L%
    df = df.drop_duplicates(subset="Metropolitan area").reset_index()
    df = df.drop(columns="index")
    
    return df

def nba_correlation():
    # YOUR CODE HERE
#     raise NotImplementedError()
    df = clean_nba_df()

    population_by_region = df["Population (2016 est.)[8]"] # pass in metropolitan area population from cities
    win_loss_by_region = df["W-L%"] # pass in win/loss ratio from nhl_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q2: Your lists must be the same length"
    assert len(population_by_region) == 28, "Q2: There should be 28 teams being analysed for NBA"

    result = stats.pearsonr(population_by_region, win_loss_by_region)
    return stats.pearsonr(population_by_region, win_loss_by_region)
```

## Question 3
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **MLB** using **2018** data.

```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

mlb_df=pd.read_csv("assets/mlb.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def clean_mlb_df():
    # load data
    mlb_df=pd.read_csv("assets/mlb.csv")
    cities=pd.read_html("assets/wikipedia_data.html")[1]
    cities=cities.iloc[:-1,[0,3,5,6,7,8]]

    # cleaning the cities dataframe
    cities["MLB"] = cities["MLB"].apply(lambda x: re.sub(r"\[.+\]", "", x))
    cities["MLB"] = cities["MLB"].replace({"Blue Jays": "Blue,Jays", 
                                           "CubsWhite Sox": "Cubs,White,Sox", 
                                           "DodgersAngels": "Dodgers,Angels", 
                                           "GiantsAthletics": "Giants,Athletics", 
                                           "YankeesMets": "Yankees,Mets",
                                           "Red Sox": "Red,Sox"})
    cities["MLB"] = cities["MLB"].apply(lambda x: x.split(","))
    cities = cities.explode("MLB")

    # cleaning the nhl_df dataframe
    mlb_df = mlb_df[mlb_df["year"] == 2018]
    mlb_df["team"] = mlb_df["team"].apply(lambda x: x.split(" ")[-1])

    # merge the dataframes
    df = pd.merge(cities, mlb_df, left_on="MLB", right_on="team")
    df = df[["Metropolitan area", "Population (2016 est.)[8]", "MLB", "team", "W", "L", "W-L%"]]
    df["Population (2016 est.)[8]"] = df["Population (2016 est.)[8]"].astype("float")
    df["W-L%"] = df["W-L%"].astype("float")

    # drop duplicated columns
    df.loc[df["Metropolitan area"] == "New York City", "W-L%"] = df.loc[df["Metropolitan area"] == "New York City", "W-L%"].mean() # mean of NY W-L%
    df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"] = df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"].mean() # mean of LA W-L%
    df.loc[df["Metropolitan area"] == "San Francisco Bay Area", "W-L%"] = df.loc[df["Metropolitan area"] == "San Francisco Bay Area", "W-L%"].maen() # mean of SF W-L%
    df.loc[df["Metropolitan area"] == "Chicago", "W-L%"] = df.loc[df["Metropolitan area"] == "Chicago", "W-L%"].mean() # mean of CH W-L%
    df.loc[df["Metropolitan area"] == "Boston", "W-L%"] = df.loc[df["Metropolitan area"] == "Boston", "W-L%"].mean() # mean of BO W-L%
    df = df.drop_duplicates(subset="Metropolitan area").reset_index()
    df = df.drop(columns="index")
    
    return df

def mlb_correlation(): 
    # YOUR CODE HERE
#     raise NotImplementedError()
    df = clean_mlb_df()

    population_by_region = df["Population (2016 est.)[8]"] # pass in metropolitan area population from cities
    win_loss_by_region = df["W-L%"] # pass in win/loss ratio from mlb_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q3: Your lists must be the same length"
    assert len(population_by_region) == 26, "Q3: There should be 26 teams being analysed for MLB"

    result = stats.pearsonr(population_by_region, win_loss_by_region)

    return stats.pearsonr(population_by_region, win_loss_by_region)
```

## Question 4
For this question, calculate the win/loss ratio's correlation with the population of the city it is in for the **NFL** using **2018** data.

```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

nfl_df=pd.read_csv("assets/nfl.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

def clean_nfl_df():
    # load data
    nfl_df=pd.read_csv("assets/nfl.csv")
    cities=pd.read_html("assets/wikipedia_data.html")[1]
    cities=cities.iloc[:-1,[0,3,5,6,7,8]]

    # cleaning the cities dataframe
    cities["NFL"] = cities["NFL"].apply(lambda x: re.sub(r"\[.+\]", "", x))
    cities["NFL"] = cities["NFL"].replace({"GiantsJets": "Giants,Jets",
                                           "RamsChargers": "Rams,Chargers",
                                           "49ersRaiders": "49ers,Raiders"
                                           })
    cities["NFL"] = cities["NFL"].apply(lambda x: x.split(","))
    cities = cities.explode("NFL")

    # cleaning the nhl_df dataframe
    nfl_df = nfl_df[nfl_df["year"] == 2018]
    nfl_df["team"] = nfl_df["team"].apply(lambda x: re.sub(r"(\*|\+)", "", x))
    nfl_df["team"] = nfl_df["team"].apply(lambda x: x.split(" ")[-1])

    # merge the dataframes
    df = pd.merge(cities, nfl_df, left_on="NFL", right_on="team")
    df = df[["Metropolitan area", "Population (2016 est.)[8]", "NFL", "team", "W", "L", "W-L%"]]
    df["Population (2016 est.)[8]"] = df["Population (2016 est.)[8]"].astype("float")
    df["W-L%"] = df["W-L%"].astype("float")

    # drop duplicated columns
    df.loc[df["Metropolitan area"] == "New York City", "W-L%"] = df.loc[df["Metropolitan area"] == "New York City", "W-L%"].mean() # mean of NY W-L%
    df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"] = df.loc[df["Metropolitan area"] == "Los Angeles", "W-L%"].mean() # mean of LA W-L%
    df.loc[df["Metropolitan area"] == "San Francisco Bay Area", "W-L%"] = df.loc[df["Metropolitan area"] == "San Francisco Bay Area", "W-L%"].mean() # mean of SF W-L%
    df = df.drop_duplicates(subset="Metropolitan area").reset_index()
    df = df.drop(columns="index")
    
    return df

def nfl_correlation(): 
    # YOUR CODE HERE
#     raise NotImplementedError()
    df = clean_nfl_df()
    
    population_by_region = df["Population (2016 est.)[8]"] # pass in metropolitan area population from cities
    win_loss_by_region = df["W-L%"] # pass in win/loss ratio from nfl_df in the same order as cities["Metropolitan area"]

    assert len(population_by_region) == len(win_loss_by_region), "Q4: Your lists must be the same length"
    assert len(population_by_region) == 29, "Q4: There should be 29 teams being analysed for NFL"

    result = stats.pearsonr(population_by_region, win_loss_by_region)

    return stats.pearsonr(population_by_region, win_loss_by_region)
```

## Question 5
In this question I would like you to explore the hypothesis that **given that an area has two sports teams in different sports, those teams will perform the same within their respective sports**. How I would like to see this explored is with a series of paired t-tests (so use [`ttest_rel`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.ttest_rel.html)) between all pairs of sports. Are there any sports where we can reject the null hypothesis? Again, average values where a sport has multiple teams in one region. Remember, you will only be including, for each sport, cities which have teams engaged in that sport, drop others as appropriate. This question is worth 20% of the grade for this assignment.

```python
import pandas as pd
import numpy as np
import scipy.stats as stats
import re

mlb_df=pd.read_csv("assets/mlb.csv")
nhl_df=pd.read_csv("assets/nhl.csv")
nba_df=pd.read_csv("assets/nba.csv")
nfl_df=pd.read_csv("assets/nfl.csv")
cities=pd.read_html("assets/wikipedia_data.html")[1]
cities=cities.iloc[:-1,[0,3,5,6,7,8]]

# clean the dataframes
def clean_dfs():
    nfl_df = clean_nfl_df()
    nba_df = clean_nba_df()
    nhl_df = clean_nhl_df()   
    mlb_df = clean_mlb_df()
    
    
    nfl_df = nfl_df[["Metropolitan area", "W-L%"]]
    nba_df = nba_df[["Metropolitan area", "W-L%"]]
    nhl_df = nhl_df[["Metropolitan area", "W-L%"]]  
    mlb_df = mlb_df[["Metropolitan area", "W-L%"]]
    
    return (nfl_df, nba_df, nhl_df, mlb_df)

def calculate_p_values(leagues):
    p_values = pd.DataFrame(columns=leagues.keys(), index=leagues.keys())
    for (league_name1, league_df1) in leagues.items():
        for (league_name2, league_df2) in leagues.items():
            merged_league = pd.merge(league_df1, league_df2, on="Metropolitan area")
            p_values.loc[league_name1, league_name2] = stats.ttest_rel(merged_league["W-L%_x"], merged_league["W-L%_y"])[1]
            
    return p_values

def sports_team_performance():
    # YOUR CODE HERE
#     raise NotImplementedError()
    (nfl_df, nba_df, nhl_df, mlb_df) = clean_dfs()
    
    leagues = {"NFL": nfl_df, "NBA": nba_df, "NHL": nhl_df, "MLB": mlb_df}
    p_values_dict = calculate_p_values(leagues)
    p_values = pd.DataFrame(p_values_dict).astype("float")
    
#     # Note: p_values is a full dataframe, so df.loc["NFL","NBA"] should be the same as df.loc["NBA","NFL"] and
#     # df.loc["NFL","NFL"] should return np.nan
#     sports = ['NFL', 'NBA', 'NHL', 'MLB']
#     p_values = pd.DataFrame({k:np.nan for k in sports}, index=sports)
    
    assert abs(p_values.loc["NBA", "NHL"] - 0.02) <= 1e-2, "The NBA-NHL p-value should be around 0.02"
    assert abs(p_values.loc["MLB", "NFL"] - 0.80) <= 1e-2, "The MLB-NFL p-value should be around 0.80"
    return p_values
```